<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 105105  -->
  <question type="stack">
    <name>
      <text>fb_unidir() unit test</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**fb_unidir() unit test**

This is a unit test question for the function [`fb_unidir`](https://github.com/mkraska/meclib/wiki/Feedback-for-unidirectional-supports). 

<p >[[input:objects]] [[validation:objects]]</p>
<p >[[input:names]] [[validation:names]] </p>
[[jsxgraph width='500px' height='400px' input-ref-objects="stateRef" input-ref-names="fbd_names" ]]
var mode  = "STACK";
const initstring = {#init#};
const centeredLabelStyle = {size:0, showInfobox:false, label:{offset:[-6,0], 
  anchorX:'left', anchorY:'middle'}};
// End of STACK header
[[include 
src="https://raw.githubusercontent.com/mkraska/meclib/main/meclib.js" /]]
[[/jsxgraph]]


Is the `names` field as expected?

[[feedback:names]]

`fb_unidir(obj, names, 3, "force \(S_3\): ")` Force \\\(S_3\\\) is not an unidirectional support

[[feedback:S3]]

`fb_unidir(obj, names, 2, "bar 2: ")` Bar 2 is not de-activated

[[feedback:fb]]

`fb_unidir(obj, names, 4, "bar 4: ")` Bar 4 has no reaction

[[feedback:fb2]]

`fb_unidir(obj, names, 5, "bar 5: ")` more than one reaction

[[feedback:fb4]]

`fb_unidir(obj, names,8, "bar 8: ")` reaction is not a force

[[feedback:fb5]]

`fb_unidir(obj, names,9, "bar 9: ")` reaction has wrong direction

[[feedback:fb8]]

`fb_unidir(obj, names,11, "bar 11: ")` correct reaction

[[feedback:fb9]]

<p hidden> spare PRTs for later tests

[[feedback:fb12]]

[[feedback:fb14]]

[[feedback:fb16]] 

</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>0</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* history
2023 04 08: cleanup local variables in fb_fbd.mac
*/
counter:1;

/* input with correct direction and name */
initdata:  [
  [ "grid", "", "", -1,10,-2,4,40],
  [ "bar", "2", [0,0], [0,2], "show"],
  [ "force", "S_3", [0,0], [0,-1], 1, "active"],
  [ "bar", "4", [1,0], [1,2], "hide"],
  [ "bar", "5", [2,0], [2,2], "hide"], 
  [ "force", "S_6", [2,0], [2,-1], 1, "active"],
  [ "moment", "M_7", [2,2], [1.5,2.5], [2.5,2.5], "active"],
  [ "bar", "8", [3,0], [2,2], "hide"],
  [ "bar", "9", [4,0], [4,2], "hide"],
  [ "force", "S_10", [4,0.5], [4.5,-0.5], 1, "active"],
  [ "bar", "11", [5,0], [5,2], "hide"],
  [ "force", "SK_12", [5,0.5], [5,-0.5], 1, "active"]
];
init: stackjson_stringify(initdata) ;

names1: [0,"show",S_3,"hide",[6,7],S_6,M_7,[7],[10],S_10,[12],SK_12];

debug: true;

stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_fbd.mac") ;

FB_9: fb_unidir(initdata, names1,9, "bar 9: ");]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>names</name>
      <type>algebraic</type>
      <tans>[]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>objects</name>
      <type>string</type>
      <tans><![CDATA[""]]></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>fb</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 2, "bar 2: ");
E: "bar 2: Nicht deaktiviert.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>StringSloppy</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_bar_name(obj, names, 12);
E: castext( "<br>{@SK_12@}: Physikalische Gr&ouml;&szlig;en immer nur mit einem Buchstabe und eventuell Index schreiben, z.B. <code>F_1</code>.");]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb12-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {@str@}<br>
E: {@E@}<br>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb14</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_bar_name(obj, names, 14);
E: castext("<br>{@A_14@}: Sinnvolle Namen für eine Stabkraft sind {@F@} (Kraft), {@N@} (Normalkraft) oder {@S@} (Stabkraft)" );]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb14-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb14-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {@str@}<br>
E: {@E@}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb16</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects) ;
[str, OK]: fb_bar_name(obj, names, 16);
E: castext("<br>Der Index 17 passt nicht zum Stabname 16." );]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb16-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb16-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {@str@}<br>
E: {@E@}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 4, "bar 4: ");
E: "bar 4: Keine Reaktion gefunden.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>StringSloppy</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 5, "bar 5: ");
E: "bar 5: Mehr als eine Reaktion gefunden.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb4-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names,8, "bar 8: ");
E: "bar 8: Die Reaktion muss eine Kraft sein.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb5-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names,9, "bar 9: ");
E: castext("bar 9: Die Reaktion \\({@S_10@}\\) hat nicht die richtige Richtung.");]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb8-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names,11, "bar 11: ");
E:  castext("bar 11: Reaktion \\({@SK_12@}\\) gefunden.");]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb9-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>names</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>flatten(names)</sans>
        <tans>flatten(names1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>names-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@names@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>names-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[Result/Expected:<br>
{@names@}<br>
{@names1@}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>S3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 3, "force \(S_3\): ");
E: "force \(S_3\): object 3 (force) is not an unidirectional support.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>S3-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>S3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>