<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 105359  -->
  <question type="stack">
    <name>
      <text>fb_vars() interactive tryout</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**fb_vars() interactive tryout**

This is a question for interactive tryout of the function [`fb_vars`](https://github.com/mkraska/meclib/wiki/fb_vars). 

**Interactive test**

Here you can try the behaviour of the feedback function. Provide `sans`, `tans` and `vars` and `label`.

<p><code>fb_vars([[input:sans]], [[input:tans]], [[input:vars]], [[input:label]])</code></p> [[validation:sans]] [[validation:tans]] [[validation:vars]] [[validation:label]] 

[[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* history
2023 04 08 reduced to interactive tryout
2023 03 14 initial version
2023 01 14 Automate the comparison by use of Answer tests (requires STACK 4.4.2);
*/

stack_unit_si_declare(true);

zero: 0 ; 
debug:true;


stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_value.mac" );



simp:false; /* to convey argumentes into fb_unit() */

E01: fb_vars(a, a);
E02: fb_unit(1.1*m, stackunits(1.101, m),0.005);
E03: fb_unit(1.1*m, stackunits(1.1001, m),0.005);
E03a: fb_unit((300*kg*m)/s^2, (300*kg*m)/s^2, 0.0);
E13: fb_unit(1, stackunits(1, m),0.005);
E14: fb_unit(-1.1*m, stackunits(1.11, m),0.005);
E15: fb_unit(1.1*m, stackunits(11, m),0.005);
E16: fb_unit(1.0*m, stackunits(11, m),0.005);
E17: fb_unit(100*m, stackunits(11, m),0.005);
E18: fb_unit(0*m, 0,0.005);
E19: fb_unit(0.1, 0.01,0.005);
E20: fb_unit(1.0, 11, 0.005);
E21: fb_unit(100, 11, 0.005);
E22: fb_unit(1, 1.001, 0.005);

lbl: castext("[[lang code='en, other']]english[[/lang]][[lang code='ru']]русский[[/lang]]");]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>label</name>
      <type>algebraic</type>
      <tans><![CDATA[""]]></tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint><![CDATA[""]]></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>sans</name>
      <type>algebraic</type>
      <tans>a+b</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>3</insertstars>
      <syntaxhint>sans</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>tans</name>
      <type>algebraic</type>
      <tans>b+c</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>3</insertstars>
      <syntaxhint>tans</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>vars</name>
      <type>algebraic</type>
      <tans>[]</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions>0.005</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@fb_vars(sans, tans,vars,label)@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_vars(sans, tans,vars,label)@}</text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>