<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 46  -->
  <question type="stack">
    <name>
      <text>QA 03 fb_unidir unit test</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**QA 03**

This is a unit test question for the function [`fb_unidir`](https://github.com/mkraska/meclib/wiki/Feedback-for-unidirectional-supports). 

<p >[[input:objects]] [[validation:objects]]</p>
<p >[[input:names]] [[validation:names]] </p>
[[jsxgraph width='500px' height='400px' input-ref-objects="stateRef" input-ref-names="fbd_names" ]]
var mode  = "STACK";
const initstring = {#init#};
const centeredLabelStyle = {size:0, showInfobox:false, label:{offset:[-6,0], 
  anchorX:'left', anchorY:'middle'}};
// End of STACK header
[[include 
src="https://raw.githubusercontent.com/mkraska/meclib/main/meclib.js" /]]
[[/jsxgraph]]


Contents of the `names` input field correct? This is pre-requisite for correct feedback.

[[feedback:names]]

**This is a debug message independent of  `debug: true`**. 

The object 3 is a force. No point to apply the function here. 

Object 11 is a fixed support, it is bi-directional, not unidirectional.

This error is independent of the student's input and must be fixed by the question author.

`fb_unidir(obj, names, 3, "3: ")` 

[[feedback:fb3]]

`fb_unidir(obj, names, 11, "11: ")` 

[[feedback:fb18]]

**Test of interactive input**.

Non-deactivated object

`fb_unidir(obj, names, 2, "2: ")` 

[[feedback:fb2]]

Object with no reaction

`fb_unidir(obj, names, 4, "4: ")` 

[[feedback:fb4]]

`fb_unidir(obj, names, 5, "5: ")`  Bar 5 has more than one reaction (one moment and one force)

[[feedback:fb5]]

`fb_unidir(obj, names, 8, "8: ")`  Bar 8 has a reaction but that is not a force

[[feedback:fb8]]

`fb_unidir(obj, names, 9, "9: ")`  Reaction force not parallel to the support

[[feedback:fb9]]

`fb_unidir(obj, names, 12, "12: ")`  Correct reaction force

[[feedback:fb12]]]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* input with correct direction and name */
initdata:  [
  [ "grid", "x", "y", -1,10,-4,4,40],
  [ "bar", "2", [0,0], [0,2], "show"],
  [ "force", "S_3", [0,0], [0,-1], 1, "active"],
  [ "bar", "4", [1,0], [1,2], "hide"],
  [ "bar", "5", [2,0], [2,2], "hide"], 
  [ "force", "S_6", [2,0], [2,-1], 1, "active"],
  [ "moment", "M_7", [2,2], [1.5,2.5], [2.5,2.5], "active"],
  [ "bar", "8", [3,0], [2,2], "hide"],
  [ "bar", "9", [4,0], [4,2], "hide"],
  [ "force", "S_10", [4,0.5], [4.5,-0.5], 1, "active"],
  [ "fix12", "11", [0, -2], 0],
  [ "bar", "12", [5,0], [5,2], "hide"],
  [ "force", "S_12", [5,0.5], [5,-0.5], 1, "active"],
  [ "bar", "14", [6,0], [6,2], "hide"],
  [ "force", "S_14", [6,2], [6,1], 1, "active"],
  [ "bar", "16", [6,2], [7,0], "hide"],
  [ "force", "S_16", [6,2], [6.5,1], 1, "active"]
];
init: stackjson_stringify(initdata) ;

names1: [0,"show",S_3,"hide",[6,7],S_6,M_7,[7],[10],S_10,"locked",[13],S_12,[15,17],S_14,[15,17],S_16] ;

debug: true;

stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_fbd.mac") ;

]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>names</name>
      <type>algebraic</type>
      <tans>[]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>objects</name>
      <type>string</type>
      <tans><![CDATA[""]]></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>fb12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 12, "12: ");
E:  ["%root","12: Reaktion \\(",["smlt","{{S}_{12}}"],"\\) gefunden."];]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb12-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}<br>
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb18</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 11, "11: ");
E: "11: object 11 (fix12) is not an unidirectional support.";
]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb18-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb18-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 2, "2: ");
E: "2: Nicht deaktiviert.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}<br>
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 3, "3: ");
E: "3: object 3 (force) is not an unidirectional support.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb3-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 4, "4: ");
E: "4: Keine Reaktion gefunden.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb4-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}<br>
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 5, "5: ");
E: "5: Mehr als eine Reaktion gefunden.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>String</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb5-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}<br>
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 8, "8: ");
E:  "8: Die Reaktion muss eine Kraft sein.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb8-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}<br>
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>fb9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[obj: stackjson_parse(objects);
[str, OK]: fb_unidir(obj, names, 9, "9: ");
E:  ["%root","9: Die Reaktion \\(",["smlt","{{S}_{10}}"],"\\) hat nicht die richtige Richtung."];]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>str</sans>
        <tans>E</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>fb9-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@str@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>fb9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>
R: {#str#}<br>
E: {#E#}<br>
[[debug /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>names</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>flatten(names)</sans>
        <tans>flatten(names1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>names-1-T</trueanswernote>
        <truefeedback format="html">
          <text>{@names@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>names-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[Result/Expected:<br>
{@names@}<br>
{@names1@}]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>